---
title: 'Topic 0 - Introductory Exercises'
description: ""
---

## Hello world!

```yaml
type: NormalExercise
key: e3c880f914
xp: 100
```

Before we dive into any complicated programming, it is important to first understand how we can view and represent objects in Python. You may think of Python as a translator that works between you and the computer. 

Instead of dealing with`1`s and `0`s, we can give Python simple commands to execute. Python allows us to view a visual representation of our data with the `print()` function.

Here's and example:
```
print("Bonjour")
```
**Note the quotation marks are required in order to print it, more on that later on**

`@instructions`
Write your first ever program!

Try and `print` the following text:
```
Hello world!
```

`@hint`
Use the `print()` function with your required text between quotation marks.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
Ex().has_output("Hello world!", pattern = False)
success_msg("Passed the first test case!")
```

---

## Python as a Calculator

```yaml
type: VideoExercise
key: 6ce7340466
xp: 50
```

`@projector_key`
25c9b7aea1a61fa9abc7d0c8fa6db8c0
